# Deeltoets2 Toegepaste Machine Learning
**Bachelor Informatiekunde**

## Tentameninformatie

**Datum:** 2023
**Tijd:** 2 uur  (+30 minuten extra tijd)


**Toegestane hulpmiddelen:** Boek als PDF op computer, Kladpapier, handgeschreven spiekbrief (1 A4, 1 kant beschreven)

## Huishoudelijke mededelingen
* Je mag tot uiterlijk 30 minuten na de begintijd starten.
* Jouw telefoon, jas en tas liggen in de kluisjes buiten de tentamenzaal.
* De eerste 30 minuten en de laatste 15 minuten mag je de zaal niet verlaten. Daarbuiten kan je als je het tentamen af  en ingeleverd hebt de zaal verlaten.
* De ondersteuning voorin de zaal kan controleren of jouw tentamen correct is ingeleverd.
* Plaats jouw collegekaart of een geldig legitimatiebewijs op de hoek van de tafel. Een surveillant komt langs om deze te controleren.
* Toiletbezoek is alleen mogelijk met toestemming van een surveillant.

## Spiekbrief

* Laat die liggen, die lever je in.

## Punten

* Elke vraag is 1 punt waard, tenzij er wat anders bijstaat.
* Bij deelvragen krijg je in de regel een evenredig deel van de punten.

## Notebook made by

__Namen__:Anoniem
__Emails__:Anoniem
__Student id__:Anoniem

## Toelichting

* De meeste opgaven worden automatisch nagekeken. Bij vrijwel alle opdrachten staan er een paar tests onder de opdracht, dit is voornamelijk om te zorgen dat je de juiste type output geeft. Dit zijn dus *NIET* alle tests, die komen er bij het graden nog bij.

## Voor het inleveren!

* Pas niet de cellen aan, vooral niet die je niet kunt editen. Copieer ook geen cellen. Dit levert problemen op bij nakijken. Je mag wel cellen toevoegen (als kladblaadje).
* Twijfel je of je per ongeluk iets hebt gewijzigd, kopieer dan bij inleveren je antwoorden naar een nieuw bestand, zodat het niet fout kan gaan.

* Zorg dat de code goed runt van boven naar beneden, verifieer dat door boven in Kernel -> Restart & Run All uit te voeren

```python
import numpy as np
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt

import sklearn
from scipy.stats import mode
from sklearn import linear_model
from sklearn.linear_model import LinearRegression
from sklearn.preprocessing import PolynomialFeatures
from sklearn import datasets
# testing
from nose.tools import assert_count_equal, assert_equal, assert_almost_equals
from numpy.testing import *
from pandas.testing import assert_frame_equal

# Please do not remove this: 
np.random.seed(31415)
```

# 1 The sigmoid (2pt)

Here is the sigmoid function $f(x)$ again:

$$f(x)= \frac{1}{1+e^x}.$$

1. What is the range and the domain of $f(x)$? 
    * Put your answer in the markdown cell below. 
2. Create a dict `D` with keys `range(-10,10)` and as values the sigmoid of the key.
3. Plot that function.
    * Put your answers in the code cell below.

**Hint:** think of `np.exp`.

<!-- grade_id: siga, points: 1, grade: true, solution: true -->
The range of f(x) is the values between -x and x. The domain of f(x) is 1.

<!-- grade_id: siga1, points: 1, grade: true, solution: true -->
```python
def sigmoid(x):
    D={x:1/(1+np.exp([-i for i in x])) for x in range(-10,10)}
    Ddf=pd.Series(D)
    return Ddf

# Plot f(x) for x in range(-10,10)
#WRITE YOUR CODE HERE
```

# 2 Logistic regression 1 (4pt)

Below we load  the penguin classification dataset, in which we want to learn to classify the `species`.


1. How many classes are there in the dataset? (variable `classes`)
2. How many instances are there in the dataset? (variable `instances`)
3. If we would do logistic regression with only the numeric variables, how many parameters (those $\theta$'s) will regression learn? Hint: write down the regression equation. (variable `parameters_numeric`)
4. And how many if we use all variables in the dataset and one-hot encode the categorical ones (variable `parameters_all`).

**Hint**: `p` is a pandas dataframe, and has all pandas functionality. E.g., `value_counts`. And you can easily select columns using method syntax, e.g, `p.species`.

```python
p=sns.load_dataset('penguins')
print(p.shape)
p.head()
```

Output:
```
(344, 7)
  species     island  bill_length_mm  bill_depth_mm  flipper_length_mm  \
0  Adelie  Torgersen            39.1           18.7              181.0   
1  Adelie  Torgersen            39.5           17.4              186.0   
2  Adelie  Torgersen            40.3           18.0              195.0   
3  Adelie  Torgersen             NaN            NaN                NaN   
4  Adelie  Torgersen            36.7           19.3              193.0   

   body_mass_g     sex  
0       3750.0    Male  
1       3800.0  Female  
2       3250.0  Female  
3          NaN     NaN  
4       3450.0  Female
```

<!-- grade_id: lr1a, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import LogisticRegression
classes=3 # change into your answer
instances=35
parameters_numeric=0
parameters_all=0

#WRITE YOUR CODE HERE

for v in [classes,instances,parameters_numeric,parameters_all]:
    print(v)
```

Output:
```
3
35
0
0
```

<!-- grade_id: lr1t, points: 4.01, grade: true, solution: false -->
```python
try:
    for v in [classes,instances,parameters_all,parameters_numeric]:
        assert isinstance(v,int)
except Exception as e: print(e)
 
   
```

# 3 Data preparation (3pt)

We look again at the penguins dataset. We call this dataframe `p`.

1. Which variables contain missing values? (variable `Vmissing`) 
2. If needed, impute the numeric variables with the mean value and the categorical with the mode (the most often occuring value).  Give the imputation values for the variablles with missing values in a dict `impute_values` of the form `{'column_name': imputation value}`.
3. For further processing simply remove the rows with missing values using `p.dropna(inplace=True)`.
4. Now Z-normalize all numeric values in p. (Z-normalization is what the sklearn's StandardScalar does). You name  the new dataframe with only the numeric columns with their  normalized values  `pz`.

<!-- grade_id: lr2a, points: 0, grade: false, solution: true -->
```python
def number_of_missing_values(dataframe) -> dict:
    {key: dataframe[key].isnull().sum() for key in dataframe.columns}
    return missing_values

Vmissing=set() # put the column names with missing data in this set.
impute_values = dict()
#WRITE YOUR CODE HERE
print(Vmissing)
print(impute_values)
```

Output:
```
set()
{}
```

<!-- grade_id: lr3a, points: 0, grade: false, solution: true -->
```python
# keep this
p.dropna(inplace=True)  # remove all rows with at least one missing value

pz=p[['bill_depth_mm', 'bill_length_mm',  'body_mass_g', 'flipper_length_mm']].copy()  
# now further work on pz to give your solution. Keep these columns in this order

#WRITE YOUR CODE HERE
pz.head()
```

Output:
```
bill_depth_mm  bill_length_mm  body_mass_g  flipper_length_mm
0           18.7            39.1       3750.0              181.0
1           17.4            39.5       3800.0              186.0
2           18.0            40.3       3250.0              195.0
4           19.3            36.7       3450.0              193.0
5           20.6            39.3       3650.0              190.0
```

<!-- grade_id: lr2t, points: 3.01, grade: true, solution: false -->
```python
try:
    assert isinstance(Vmissing,set)
    assert Vmissing.issubset(set(p.columns))
    assert isinstance(impute_values,dict)
    assert pz.shape==(333, 4)
    assert_array_equal(list(pz.columns), ['bill_depth_mm', 'bill_length_mm',  'body_mass_g', 'flipper_length_mm'])
except Exception as e: print(e)

```

# 4 Decision trees (4pt)

Now we will predict the species in the penguin dataset.

1. First run the cell below to make the X and y sets.
2. Split X and y into train and test sets using a 2/3, 1/3 split using `train_test_split` with `random_state=0`.
3. Then use the `DecisionTreeClassifier(random_state=0)` to train a classifier, and compute its accuracy in variable `accuracy_testset`.
4. Then create the report with all precision, recall and F1  values for each class
5. Finally make a confusion matrix and print it nicely with `sns.heatmap`.

```python
X= pz.values
p.dropna(inplace=True) 
y= p.species.astype('category').cat.codes.values  # here we just replace the species names with integers
# check: X and y have the same number of rows, and both no missing values

X.shape,y.shape, np.isnan(X).sum(), np.isnan(y).sum()
```

Output:
```
((333, 4), (333,), 0, 0)
```

<!-- grade_id: ml1a, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split
Xtrain, Xtest, ytrain, ytest= = train_test_split(X, y, random_state=0) # change to your solution
#WRITE YOUR CODE HERE
Xtrain.shape, Xtest.shape, ytrain.shape, ytest.shape 
```

<!-- grade_id: ml1b, points: 0, grade: false, solution: true -->
```python
from sklearn.tree import DecisionTreeClassifier

accuracy_testset=0 # change to your (hopefully better) solution
#WRITE YOUR CODE HERE
accuracy_testset 
```

<!-- grade_id: ml1t, points: 2.01, grade: true, solution: false -->
```python
try: 
    assert 0<= accuracy_testset <= 1
except Exception as e: print(e)
    
```

# the report with all precision, recall and F1  values for each class

Make it in the cell below

<!-- grade_id: ml1c, points: 1, grade: true, solution: true -->
```python
from sklearn.metrics import classification_report
classification_report()
#WRITE YOUR CODE HERE
```

# The confusion matrix nicely  printed with `sns.heatmap`

* in thecell below

<!-- grade_id: ml1d, points: 1, grade: true, solution: true -->
```python
from sklearn.metrics import confusion_matrix
confusion_matrix 

#WRITE YOUR CODE HERE
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Input In [50], in <cell line: 2>()
      1 #WRITE YOUR CODE HERE
----> 2 from metrics import confusion_matrix

ModuleNotFoundError: No module named 'metrics'
```